```python
# Initialize Otter
import otter
grader = otter.Notebook("control-structures.ipynb")
```

# 3 Kontrollstrukturen

**Computational Thinking WS2022/21** 

*Prof. Dr.-Ing. Martin Hobelsberger, Dr. Benedikt Zönnchen*

[CT-Buch Kapitel 17. Kontrollstukturen](https://bzoennchen.github.io/ct-book/chapters/03-7/intro.html) 

***

Es gibt zwei wesentliche Kontrollstrukturen:

1. **Fallunterscheidungen:** Führe entweder Codeabschnitt A oder B oder C ... aus
2. **Schleifen:** Führe Codeabschnitt A wiederholt aus

## 3.1 Fallunterscheidungen

[CT-Buch Kapitel 17.1 Fallunterscheidungen](https://bzoennchen.github.io/ct-book/chapters/03-7/1-cases.html) 

Über sogenannte `if`-*Statements* kann der Programmablauf gesteuert werden.
Dabei wird eine *Befehlsfolge* an eine **Bedingungen** geknüpft: 

*Computer, nur wenn Bedingung A zutrifft führe Codezeilen 4-6 aus!*

```python
if Bedingung A:
    Befehl 1
    Befehl 2
    ...
elif Bedingung B: #und nicht Bedingung A
    Befehl 3
    Befehl 4
    ...
else: # weder Bedingung A noch Bedingung B
    Befehl 5
    Befehl 6
    ...
```

<span style='color:#C70039'>**Wichtig: Die Einrückung und der ``:`` spielen DIE entscheidende Rolle!**</span>

```python
# if

```

```python
# if-else

```

```python
# if-elif-elif

```

```python
# if-elif-elif-else

```

```python
# if if if-else

```

```python
# if inside if

```

#### Anmerkung:
* Die Doppelpunkte nicht vergessen!
* Es gibt **IMMER** eine else-Verzweigung (default)
* Die Einrückung ist sehr wichtig!!

***
***Aufgabe 3.1*** Schreiben Sie ein Programm um *Schere-Stein-Papier* zu spielen.

```python
scissors = 'Schere'
stone = 'Stein'
paper = 'Papier'

player1 = input(f'Spieler 1: {scissors}, {stone}, {paper}? ')
player2 = input(f'Spieler 2: {scissors}, {stone}, {paper}? ')

...
```

***

## 3.2 Schleifen

[CT-Buch Kapitel 17.2 Schleifen](https://bzoennchen.github.io/ct-book/chapters/03-7/2-loops.html) 

Erst durch die Wiederholen werden Computer zu mächtigen Rechenmaschine und *Schleifen* verkörpern diese Wiederholung wie kaum eine Struktur. Dabei gibt es in ``Python`` zwei Typen von Schleifen:

1. ``for``-Schleifen
2. ``while``-Schleifen

### 3.2.1 Die For-Schleife

Mit der ``for``-Schleife können Sie eine Folge an Befehlen (*Statements*) für eine, **vor dem Eintritt in die Wiederholung bekannte Anzahl**, ausführen.
Sie können die ``for``-Schleife jedoch abhängig von einer Bedingung frühzeitig verlassen/abbrechen.

Dabei *iteriert* die ``for``-Schleife über eine *iterierbare* Ansammlung (*Collection*) an Objekten (z.B. Zahlen, Zeichen, etc.).
Der Codeblock der Schleife erlangt für jede Wiederholung Zugriff auf das nächste Element der Ansammlung.
In welcher Reihenfolge, hängt von der Datenstruktur ab, welche diese Ansammlung repräsentiert. 

```python
for element in elements: #Für jedes Element in der Menge Elemente
    # führe folgenden Codeblock abhängig von element aus
    pass
```

```python
# for liste

```

```python
# for range

```

***
***Aufgabe 3.2*** Schreiben Sie ein Programm was Ihnen die Summe ``sum_squares`` der Quadratzahlen von ``1`` bis ``n`` berechnet.

```python
n = 12
```

```python
sum_squares = 0
...
sum_squares
```

```python
grader.check("q2")
```

***

```python
# for _
```

```python
names = ['Sarah', 'Sebastian', 'Bene', 'Anna', 'Martin', 'Sabine']

# for
```

```python
# for chars

```

Mit ``break`` können wir eine Schleife frühzeitig verlassen.

```python
# break

for char in 'Bene':    
    if char == 'n':
        break
    print(char)
```

Mit ``continue`` können wir die Ausführung des aktuellen Wiederholung abbrechen.

```python
# continue

for char in 'Bene':
    if char == 'n':
        continue
    print(char)
```

***
***Aufgabe 3.3*** Schreiben Sie eine Funktion ``is_prime(n)``, die testet ob ``n`` eine Primzahl ist oder nicht.

```python
def is_prime(n):
    ...
```

```python
print(is_prime(2)) # True
print(is_prime(3)) # True
print(is_prime(23)) # True
print(is_prime(55)) # False
print(is_prime(1)) # False
```

```python
grader.check("q3")
```

***

### 3.2.2 Die While-Schleife

Mit der ``while``-Schleife können Sie eine Folge an Befehlen (*Statements*) solange ausführen, wie eine festgelegt Bedingung zutrifft.
Sie können die ``while``-Schleife jedoch abhängig von einer anderen Bedingung auch frühzeitig verlassen/abbrechen.

Sie müssen nicht wissen wann die Bedingung erfüllt ist, doch sollten sie irgendwann erfüllt sein!
Anderfalls endet die Wiederholung niemals.

Dies macht die ``while``-Schleife mächtiger und zugleich fehleranfälliger als die ``for``-Schleife. 

```python
while condtion: #Solange Bedingung erfüllt
    # führe folgenden Codeblock aus
    pass
```

Jede ``for``-Schleife lässt sich in eine ``while``-Schleife umwandeln, umgekehrt gilt dies nicht.

```python
# for into while

```

```python
# for 

```

**Faustregel:** Verwenden Sie die ``for``-Schleife wenn immer Sie sich eignet.

***
***Aufgabe 3.4*** Berechnen Sie die größte Primzahl ``prime``, die kleiner als ``1000`` ist.

```python
# solve it

```

***Aufgabe 3.5*** Addieren Sie alle Primzahlen von ``2`` aufsteigend auf bis die Summe durch ``42`` teilbar ist.

```python
prime_sum = ...
...
print(prime_sum)
print(prime_sum % 42 == 0)
```

***

```python
# endless game

    
```

```python

```

## Abgabe

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q2 (points: None)
```
>>> sum_squares == sum(map(lambda x: x**2, range(n)))
True
```


Tests for q3 (points: None)
```
>>> is_prime(1) == False
True
```
```
>>> is_prime(2) == True
True
```
```
>>> is_prime(71) == True
True
```
```
>>> is_prime(24) == False
True
```
```
>>> is_prime(39) == False
True
```
