```python
# Initialize Otter
import otter
grader = otter.Notebook("expressions.ipynb")
```

# 2 Datentypen und Befehle

**Computational Thinking WS2022/21** 

*Prof. Dr.-Ing. Martin Hobelsberger, Dr. Benedikt Zönnchen*

+ [CT-Buch Kapitel 13.2 Ausdrücke](https://bzoennchen.github.io/ct-book/chapters/03-3/2-expressions.html) 
+ Den Anfang von [CT-Buch Kapitel 14. Datentypen (Grundlagen)](https://bzoennchen.github.io/ct-book/chapters/03-4/intro.html)

***



Jedes Programm, bzw. jeder Algorithmus besteht aus einer Folge von *Befehlen/Aussagen* (engl. *Statements*). Diese *Befehle* beschreiben wie Daten (die Eingabe) verarbeitet werden sollen. Ein Kochrezept besteht beispielweise aus Befehlen, die beschreiben wie ein Gericht zubereitet werden soll. Die Initialisierung und Zuweisung einer *Variablen*

```python
# Statements


```

ist ein *Befehl* (engl. *Statement*). 

Befehle die einen Wert ergeben bzw. zurückliefern, bezeichnen wir als *Ausdruck* (engl. *Expression*).
Die Multiplikation 

```python
# Expressions


```

ist beispielsweise ein *Ausdruck* der zwei Dezimalzahlen multipliziert. Der *Ausdruck* besteht aus dem Symbol ``*`` und zwei nummerischen Werten (Zahlen)

``[Zahl] * [Zahl]``

welche wiederum *Ausdrücke* sind.

## 2.1 Syntax und Semantik

Welche *Befehle* bzw. *Ausdrücke* valide sind, gibt die **Syntax** der verwendeten Programmiersprache (hier ``Python``) vor. Welchen *Effekt* bzw. *Bedeutung* die *Befehle* bzw. *Ausdrücke* haben, gibt deren **Semantik** vor.

Die Multiplikation wird durch den Computer, genauer die CPU berechnet. ``3 * 5`` ergibt ``15``, d.h. die Semantik des Ausdrucks ``3 * 5`` ist ``15``. 

Die Zeichenfolge muss **syntatisch** korrekt sein, damit diese auch als Ausdruck vom Computer (bzw. *Interpreter*) verstanden wird. Die Zeichenfolge

```python
# Syntax Error

```

ist kein *syntaktisch* korrekter ``Python``-Befehl, was Ihnen die Fehlermeldung auch zu verstehen gibt. Die gewählte Programmiersprache bestimmt (durch ihre Grammatik) welche Zeichenfolge **syntaktisch** korrekt sind. Bereits kleiner Änderungen an der **Syntax** können zu einer neuen *Bedeutung (Semantik)* führen. 

In ``Python`` ist die Zeichenkette

```python
3 ** 5
```

ein *syntaktisch* valider Ausdruck und *bedeutet* $3^5$ was $243$ ergibt.

Der Computer ist äußerst primitiv und benötigt exakte Anweisungen.
Sie können beispielweise eine Multiplikation nicht wie folgt beschreiben:

```python
3 multiply 5
```

da die Multiplikation nur mit dem ``*``-Operator ausgelöst werden kann.

```python
3 * 5
```

Bezüglich des Programmierens sind die Konzepte **Syntax** (Grammatik) und **Semantik** sowie das **Lernen durch Praktizieren** ähnlich zu den natürlichen Fremdsprachen. Allerdings sind die Regeln sehr viel einfacher, da es keinerlei Ausnahmen oder Sonderfälle gibt!

Ist die **Syntax** Ihres Programms korrekt bedeutet dies noch nicht, dass Sie auch das berechnen, was sie berechnen wollen. Es bedeutet lediglich, dass Ihr Code in der Sprache liegt.

## 2.2 Zusammensetzen von Aussagen

*Aussagen* und *Ausdrücke* können sich aus weiteren *Aussagen* und *Ausdrücken* zusammensetzen. Der ``+``-Operator erwarten zwei Ausdrücke deren Semantik eine Zahl ergibt: ``[Ausdruck (Zahl)] + [Ausdruck (Zahl)]`` gleichzeitig wird dieser Ausdruck selbst zur Zahl *ausgewertet*.

```python
# Expression

```

ist ein Ausdruck und

```python
# Expression

```

ebenfalls und

```python
# Expression

```

auch.

## 2.3 Datentypen (Grundlagen)

Ob ein *Befehl* oder *Ausdruck* syntaktisch korrekt ist und was dieser semantisch ergibt, hängt von den Datentypen der zu verarbeitenden Werte ab!
Addieren wir zwei ganze Zahlen (Datentyp: ``int``) erhalten wir ein anderes Ergebnis als wenn wir zwei Zeichen (Datentyp: ``str``) "Addieren".

```python
# Addition

```

```python
# Konkatination

```

Im Speicher des Rechners befinden sich ausschließlich Bits. Was diese Bits **repräsentieren** hängt vom Kontext ab. Z.B. interpretiert die Recheneinheit, welche ganze Zahlen addiert, die Bits als Binärzahlen.

Datentypen einer Programmiersprache erlauben es uns festzulegen **wie ein bestimmter Wert im Arbeitsspeicher** interpretiert werden soll.

Weisen wir der Variablen ``char`` den Wert ``'a'`` zu

```python
char = 'a'
char
```

Output:
```
'a'
```

so wird im Speicher irgendwo der **Wert** als binärer ASCII-Code stehen:

$$01100001_2$$

Würde man diese Bitfolge als ganze Zahl interpretieren wäre dies gleich

$$2^6 + 2^5 + 2^0 = 97_{10}$$

Warum gibt uns der Interpreter aber ``'a'`` und nicht ``97`` aus? Und weshalb kommt es bei folgender Addition

```python
# Failed addition str + int

```

zu einem Fehler? Die kurze Antwort lautet: Wegen der Datentypen der Variablen ``x`` bzw. ``char``.

```python
# type of x and char

```

Der Datentyp der Variablen ``char`` ist ``str`` (Zeichenkette). Diese Information erhält der Interpreter und interpretiert deshalb die Bitfolge im Speicher als Zeichenkette.
Die ``+``-Operation ist für die Kombination von Datentypen ``int`` und ``str`` nicht definiert.
Deshalb kommt es zu einem Fehler.

```python

```

### 2.3.1 Atomare Datentypen

```python
# bool, int, float
```

```python

```

### 2.3.2 Zusammengesetzte Datentypen (Datenstrukturen)

```python
# str, list
```

```python

```

```python

```

## 2.4 Arithmetische Operatoren

Die Multiplikation ``*`` wie auch die Potenz ``**`` bezeichnen wir als *arithmetische Operatoren*, da sie numerische Werte (Zahlen) verarbeiten.
Es gibt jedoch noch eine ganze Reihe von weiteren *arithmetische Operatoren*:

| Operator |     Beschreibung     |  Beispiel  |                     Bedeutung                        |
| :------: | :------------------: | :--------: | :--------------------------------------------------: |
|   `+`    |       Addition       | ``3 + 4``  |                      $3 + 4$                         |
|   `-`    |     Subtraktion      | ``3 - 4``  |                      $3 - 4$                         |
|   `*`    |    Multiplikation    | ``3 * 4``  |                    $3 \cdot 4$                       |
|   `/`    |       Division       | ``3 / 4``  |                      $3 / 4$                         |
|   `**`   |     Potenzierung     |  ``3**4``  |                       $3^4$                          |
|   `//`   | ganzzahlige Division | ``3 // 4`` |         $\left \lfloor{3/4}\right \rfloor$           |
|   `%`    |        Modulo        | ``10 % 4`` | $$10 - (4 \cdot \left \lfloor{10/4}\right \rfloor)$$ |

Jeder dieser Operatoren ``op`` erwartet zwei Zahlen, eine links und eine rechts vom jeweiligen Operator.

Die Bedeutung der Modulo-Operation sieht kompliziert aus doch bedeutet dies schlicht, dass ``10 % 4`` der ganzzahlige Rest der *Restwertdivision* ist.

Die ganzzahlige Division rundet das Ergebnis der Division auf die nächst kleinere ganze Zahl (Integer).

```python
-2 // 3
```

Output:
```
-1
```

```python
-2 / 3
```

Output:
```
-0.6666666666666666
```

```python
2 // 3
```

Output:
```
0
```

```python
2 / 3
```

Output:
```
0.6666666666666666
```

```python
14 % 5
```

Output:
```
4
```

```python
14 - (14 // 5)*5
```

Output:
```
4
```

```python
x = 12314
y = 7

x % y == x - (x // y) * y
```

Output:
```
True
```

Die Addition und Subtraktion von Fließkommazahlen kann zu merkwürdigen Ergebnissen führen. Dazu später mehr:

```python
0.1 + 0.2
```

Output:
```
0.30000000000000004
```

```python
0.1 + 0.2 == 0.3 # False! Oo?
```

Output:
```
False
```

```python
0.1 + 0.4
```

Output:
```
0.5
```

```python
0.1 + 0.4 == 0.5 # True
```

Output:
```
True
```

***
***Aufgabe 2.1*** Berechnen Sie folgenden Ausdruck

$$10^6 - 10^{-10}$$

und 

$$10^6 - 10^{-11}$$

und weisen Sie das erste Ergebnis der Variable ``x`` und das zweite Ergebnis der Variable ``y`` zu.
Lassen Sie sich ``x`` und ``y`` ausgeben. Was beobachten Sie?

_Ersetzen Sie diesen Text durch Ihre Antwort._

```python
x = ...
x
```

```python
y = ...
y
```

```python
grader.check("q1")
```

***Aufgabe 2.2*** Berechnen Sie die Wurzel aus $2$ und weisen Sie das Ergebnis der Variable ``z`` zu. Repräsentiert der Wert der Variable wirklich $\sqrt{2}$? Warum bzw. warum nicht?

_Ersetzen Sie diesen Text durch Ihre Antwort._

```python
z = ...
```

```python
grader.check("q2")
```

***

## 2.5 Vergleichsoperatoren

Werte/Objekte können über Vergleichsoperatoren miteinander verglichen werden. Das Ergebnis ist ein Boolscher Wert, d.h. ein Wahrheitswert (wahr ``True`` oder falsch ``False``).

| Operator  | Beschreibung                      |
| :-------: | :-------------------------------- |
| `x == y ` | ist `x` gleich `y`? (Gleichheit)              |
| `x != y`  | ist `x` ungleich `y`?             |
|  `x > y`  | ist `x` größer als `y`?           |
| `x >= y`  | ist `x` größer oder gleich `y`?   |
|  `x < y`  | ist `x` kleiner `y`?              |
| `x <= y`  | ist `x` kleiner gleich `y`?       |
| `x is y`  | ist `x` das selbe Objekt wie `y`? (Identität) |

Erneut ist ``Python`` hier ein wenig speziell indem es die mathematische Schreibweise $0 < x < 5$ anstatt $0 < x \ \land \ x < 5$ zulässt.
Dies erhöht die Lesbarkeit, da wir solche Verkettungen von Vergleichsoperatoren gewohnt sind.

Den Unterschied zwischen ``is`` (Identität) und ``==`` (Gleichheit) werden wir später erläutern, es sei allerdings hier schon erwähnt, dass ``x is y`` wahr (``True``) ist wenn die *Variable* ``x`` auf den gleiche Speicher verweißt wie ``y``. 
``x == y`` ist hingegen auch dann wahr (``True``) falls ``x`` und ``y`` auf die gleichen **Werte** im Speicher an unterschiedlichen Speicherbereichen verweisen.

Zwei 50 Cent Münzen sind gleich aber nicht identisch.

```python
# int Vergleiche

```

```python
# float Vergleiche 0.1 + 0.3, epsilon

```

```python

```

Der Ausdruck

```python
5 < 7 < 10
```

hat die gleiche *Bedeutung* wie der Ausdruck

```python
5 < 7 and 7 < 10
```

*Vergleichsoperatoren* können auch auf anderen Datentypen als numerische Werte (ganze Zahlen ``int``, Fließkommazahlen ``float``) definiert sein.
So können wir in ``Python`` auch Zeichenketten ``str`` mit den Vergleichsoperatoren lexikographisch vergleichen:

```python
# str Vergkeiche

```

***
***Aufgabe 2.3*** In der folgenden Zelle berechnen wir

$$\sqrt{2} + 2 - \sqrt{2} - 2 = 0$$

jedoch erhalten wir nicht $0$ sondern eine sehr kleine negative Zahl. Schreiben Sie einen Ausdruck der prüft ob ``x`` fast gleich ``0`` ist.

```python
x = 2**(0.5) + 2 - 2**(0.5) - 2
x
```

```python
epsilon = ...
...
```

***

## 2.6 Logische (Boolsche) Operatoren

Der *Ausdruck* 

```python
5 < 7 and 7 < 10
```

besteht aus den *Ausdrücken* ``5 < 7`` und ``7 < 10`` die dem logischen ``and``-Operator verknüpft werden.
Dieser erwartet auf der linken und rechte Seite jeweils einen Wahrheitswert (*boolschen Ausdruck*). *Vergleichsoperatoren* liefern sind *boolsche Ausdrücke*.
Seine Semantik ist den

$$5 < 7 \land 7 < 10$$

Der ``and``-Operator, mit ``exp1 and exp2``, ergibt genau dann ``True`` wenn die Auswertung von ``exp1`` und ``exp2`` jeweils ``True`` ergeben.

```python
exp1 = True
exp2 = 5 < 6
exp1 and exp2
```

Es gibt drei logische Operatoren:

| Operator  | Beschreibung                                               |
| :-------: | :--------------------------------------------------------- |
|  `not x`  | ist `True` genau dann wenn `x == False`.                   |
| `x and y` | ist `True` genau dann wenn `x == True` und `y == True`.    |
| `x or y`  | ist `True` genau dann wenn ``x == True`` oder `y == True`. |

```python
# Beispiele

```

## 2.7 Bitoperatoren

Der Vollständigkeit listen wir auch noch die *Bitoperatoren* auf.
Diese sind dazu vorgesehen um den Wert in *Binärdarstellung* zu manipulieren.

Jeder Wert egal ob Zahl, Zeichen, Bild, Ton wir als Binärcode im Speicher abgelegt.
*Bitoperatoren* nehmen diesen Binärwert und verarbeiten bzw. kombinieren ihn genau wie die Addition zwei Zahlen in der Dezimalschreibweise verarbeitet.

Dabei wird jedes Bit des einen Werts mit dem *Bit* des anderen Werts kombiniert.
Zum Beispiel ``5 & 4`` führt eine führ jedes Bit die ``and`` Operation aus.
Das nennen wir Verundung.

```python
5 & 4
```

führt eine Verundung der Binärzahlen $5_{10} = 101_2$ mit $4_{10} = 100_2$ durch und ergibt demnach $100_2$, was wiederum gleich $4_{10}$ ist.

| Operator | Beschreibung                                       | Beispiel | Ergebnis |
| :------: | :------------------------------------------------- | :------- | :------- |
| `x & y`  | Verundung von `x` mit `y`                          | `10 & 3` | `2`      |
| `x \| y`  | Veroderung von `x` mit `y`                         | `10 \| 3` | `11`     |
| `x ^ y`  | exklusive Veroderung von `x`  mit `y`              | `10 ^ 3` | `9`      |
| `x << y` | Bitverschiebung von `x` um `y` Stellen nach links  | `8 << 3` | `64`     |
| `x >> y` | Bitverschiebung von `x` um `y` Stellen nach rechts | `8 >> 2` | `2`      |

Weshalb ist ``10 ^ 3`` gleich ``9``?
Nun $10_{10} = 01010_2$ und $3_{10} = 00011_2$.
Das exklusive oder bedeutet gesprochen **entweder oder**, d.h. ein Bit wird zur 1 genau dann wenn das Bit der einen Zahl gleich 1 und das Bit der anderen Zahl gleich 0 ist oder genau anders herum.
Dies ergibt $01001_2 = 9_{10}$.

Für ganze Zahlen entspricht die Bitverschiebung nach rechts um ein Bit der Multiplikation mit 2.
Die Verschiebung nach rechts um ein Bit hingegen der *ganzzahligen Division* durch 2.
Deshalb ist `8 << 3` gleich $8 \cdot 2 \cdot 2 \cdot 2 = 8 \cdot 2^3 = 64$ und ``8 >> 2`` gleich $\left \lfloor{8 \cdot 2^{-2}}\right \rfloor = 2$.

***
***Aufgabe 2.4*** Berechnen Sie $x \leftarrow 2^{10}$ unter Verwendung der Bitverschiebung.

```python
x = ...
```

```python
grader.check("q4")
```

***Aufgabe 2.5*** Angenommen die Variablen ``x`` und ``y`` haben jeweils eine ganze Zahl als **Wert**, sodass beide dieser Zahlen in der Binärdarstellung an keiner Stelle ein gleiches Bit besitzten. Zum Beispiel:

$$x \leftarrow 5_{10} = 0101_2 \text{ und } y \leftarrow 10_{10} = 1010_2.$$

Mit welcher Bitoperation könnten Sie die Addition ``x + y`` berechnen? Testen Sie diesen Sachverhalt.

_Ersetzen Sie diesen Text durch Ihre Antwort._

```python
x = ...
y = ...
...
```

## Abgabe

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> x == 999999.9999999999
```
```
>>> y == 1000000.0
```


Tests for q2 (points: None)
```
>>> z == 1.4142135623730951
```


Tests for q4 (points: None)
```
>>> x == 2**10
```
